---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.

`@instructions`
You are at the bottom of everything. `asdf` is a variable.

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 45e5490216
xp: 50
```



`@possible_answers`
adf
adsf
adsasf

`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```
